## Before submitting
1. Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

2. Make sure that no assertions fail or exceptions occur, otherwise points will be subtracted.

3. Please submit only the `*.ipynb` file.

4. Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". Edit only between `YOUR CODE HERE` and `END YOUR CODE`.

5. Make sure to use Python 3, not Python 2.

Fill your group name and collaborators below:

```python
GROUPNAME = ""
COLLABORATORS = ""
```

---

# Exercise Sheet 1: Python Basics

This first  exercise sheet tests the basic functionalities of the Python programming language in the context of a simple prediction task. We consider the problem of predicting health risk of subjects from personal data and habits. We first use for this task a decision tree

![](tree.png)

adapted from the webpage http://www.refactorthis.net/post/2013/04/10/Machine-Learning-tutorial-How-to-create-a-decision-tree-in-RapidMiner-using-the-Titanic-passenger-data-set.aspx. For this exercise sheet, you are required to use only pure Python, and to not import any module, including numpy. In exercise sheet 2, the nearest neighbor part of this exercise sheet will be revisited with numpy.

## Classifying a single instance (15 P)

* Create a function that takes as input a tuple containing values for attributes (smoker,age,diet), and computes the output of the decision tree. Should return `"less"` or `"more"`.
* Test your function on the tuple `('yes', 31, 'good')`,

<!-- grade_id: cell-b66d7278bc313c94, points: 0, grade: false, solution: true -->
```python
def decision(x):
    # >>>>> YOUR CODE HERE
    if len(x) != 3:
        return 'wrong input'
    if x[0] == 'yes':
        if x[1] < 29.5:
            return 'less'
        else:
            return 'more'
    else:
        if x[2] == 'good':
            return 'less'
        else:
            return 'more'
    # <<<<< END YOUR CODE
```

<!-- grade_id: cell-c31b80471db3132f, points: 15, grade: true, solution: false -->
```python
# Test
x = ('yes', 31, 'good')
assert decision(x) == 'more'

```

```python

```

## Reading a dataset from a text file (10 P)

The file `health-test.txt` contains several fictious records of personal data and habits.

* Read the file automatically using the methods introduced during the lecture.
* Represent the dataset as a list of tuples. Make sure that the tuples have the same format as above, e.g. `('yes', 31, 'good')`.
* Make extra note of the datatype of each element

<!-- grade_id: cell-c1a8bc4c0e4ccb26, points: 0, grade: false, solution: true -->
```python
def gettest():
    # >>>>> YOUR CODE HERE
    f = open("health-test.txt", "r")
    lines = f.readlines()
    dataset = []
    for line in lines:
        #get rid of the '\n' character
        line = line[:-1]
        #split the line 
        temp_tuple = line.split(",")
        #cast the second element
        temp_tuple[1] = float(temp_tuple[1])
        #cast to tuple
        temp_tuple = tuple(temp_tuple)
        dataset.append(temp_tuple)
    return dataset
    # <<<<< END YOUR CODE
```

<!-- grade_id: cell-4e1f7ad1e66b3121, points: 10, grade: true, solution: false -->
```python

```

## Applying the decision tree to the dataset (15 P)

* Apply the decision tree to all points in the dataset, and return the ratio of them that are classified as "more".

<!-- grade_id: cell-6703ef98e2b5c93b, points: 0, grade: false, solution: true -->
```python
def evaluate_testset():
    # >>>>> YOUR CODE HERE
    dataset = gettest()
    total_of_tuples = len(dataset)
    more_classified_tuples = 0
    for i in dataset:
        if decision(i) == 'more':
            more_classified_tuples = more_classified_tuples +1
    return more_classified_tuples / total_of_tuples
    # <<<<< END YOUR CODE
```

<!-- grade_id: cell-c13a0b23c9faba52, points: 15, grade: true, solution: false -->
```python

```

## Learning from examples (10 P)

Suppose that instead of relying on a fixed decision tree, we would like to use a data-driven approach where data points are classified based on a set of training observations manually labeled by experts. Such labeled dataset is available in the file `health-train.txt`. The first three columns have the same meaning than for `health-test.txt`, and the last column corresponds to the labels.

* Write a procedure that reads this file and converts it into a list of pairs. The first element of each pair is a triplet of attributes, and the second element is the label.

<!-- grade_id: cell-fc38ed11fee6fbeb, points: 0, grade: false, solution: true -->
```python
def gettrain():
    # >>>>> YOUR CODE HERE
    f = open("health-train.txt", "r")
    lines = f.readlines()
    dataset = []
    for line in lines:
        #get rid of the '\n' character
        line = line[:-1]
        #split the line 
        temp_tuple = line.split(",")
        #cast the second element
        temp_tuple[1] = float(temp_tuple[1])
        #Now we have a 4 elemental list, so we pop the last element and store it
        label = temp_tuple.pop(3)
        #cast to tuple
        temp_tuple = tuple(temp_tuple)
        #append the pair
        dataset.append((temp_tuple, label))
    return dataset
    # <<<<< END YOUR CODE
```

<!-- grade_id: cell-a3d593f232e0403a, points: 10, grade: true, solution: false -->
```python

```

## Nearest neighbor classifier (25 P)

We consider the nearest neighbor algorithm that classifies test points following the label of the nearest neighbor in the training data. For this, we need to define a distance function between data points. We define it to be

`d(a, b) = (a[0] != b[0]) + ((a[1] - b[1]) / 50.0) ** 2 + (a[2] != b[2])`

where `a` and `b` are two tuples corrsponding to the attributes of two data points.

* Write a function that retrieves for a test point the nearest neighbor in the training set, and classifies the test point accordingly.
* Test your function on the tuple `('yes', 31, 'good')`

<!-- grade_id: cell-671ea24ec8a11241, points: 0, grade: false, solution: true -->
```python
def distance_neighbor(a,b):
    return (a[0] != b[0]) + ((a[1] - b[1]) / 50.0) ** 2 + (a[2] != b[2])

def neighbor(x, trainset):
    # >>>>> YOUR CODE HERE
    #we define a tuple, that gets the distance between x and a point in trainset, and the tuple in the trainset
    #we initalized it with infinity and None
    min_distance_tuple = (float("inf"),None)
    #We iterate over all the trainset, and everytime we find a smaller distance we update our tuple
    for i in trainset:
        temp_distance = distance_neighbor(x, i[0])
        if temp_distance < min_distance_tuple[0]:
            min_distance_tuple = (temp_distance,i)
    
    #we just want the label from the tuple so
    return min_distance_tuple[1][1]
    # Beeing min_distance_tuple[1] a tuple in the format ('yes', 33.0, 'good'), 'more')
    # And from that tuple we want the label so the [1] element
    # <<<<< END YOUR CODE
```

<!-- grade_id: cell-a36122337853f195, points: 15, grade: true, solution: false -->
```python
# Test
x = ('yes', 31, 'good')
assert neighbor(x, gettrain()) == "more"

```

* Apply both the decision tree and nearest neighbor classifiers on the test set, and return the list of data point(s) for which the two classifiers disagree, and with which probability it happens.
* A data point should look like above, e.g. `('yes', 31, 'good')`.

<!-- grade_id: cell-8dbf7da153f3d797, points: 0, grade: false, solution: true -->
```python
def compare():
    # >>>>> YOUR CODE HERE
    dataset = gettrain()
    Xdisagree = []
    for i in dataset:
        if decision(i[0]) != neighbor(i[0], dataset):
            Xdisagree.append(i[0])
    probability = float(len(Xdisagree)/len(dataset))
    # <<<<< END YOUR CODE
    return Xdisagree, probability
```

<!-- grade_id: cell-3b55f7e89ad4dfeb, points: 10, grade: true, solution: false -->
```python
Xdisagree, probability = compare()
assert type(Xdisagree) == list
print(Xdisagree)
assert probability >= 0.0 and probability <= 1.0
```

Output:
```
[('yes', 54.0, 'good')]
```

One problem of simple nearest neighbors is that one needs to compare the point to predict to all data points in the training set. This can be slow for datasets of thousands of points or more. Alternatively, some classifiers train a model first, and then use it to classify the data.

## Nearest mean classifier (25 P)

We consider one such trainable model, which operates in two steps:

(1) Compute the average point for each class, (2) classify new points to be of the class whose average point is nearest to the point to predict.

For this classifier, we convert the attributes smoker and diet to real values (for smoker: yes=1.0 and no=0.0, and for diet: good=0.0 and poor=1.0), and use the modified distance function:

`d(a,b) = (a[0] - b[0]) ** 2 + ((a[1] - b[1]) / 50.0) ** 2 + (a[2] - b[2]) ** 2`

We adopt an object-oriented approach for building this classifier.

* Implement the methods `train` and `predict` of the class `NearestMeanClassifier`.

<!-- grade_id: cell-e0b339bfd0fcc16c, points: 0, grade: false, solution: true -->
```python
def distance(a,b):
    return (a[0] - b[0]) ** 2 + ((a[1] - b[1]) / 50.0) ** 2 + (a[2] - b[2]) ** 2

#we have data in form from ((str,float,str),str) , we want ((float,float,float),str) in order to compute the mean
def convert_data(dataset):
    converted_data_set = []
    for i in dataset:
        converted_smoker = 1.0 if i[0][0] == "yes" else 0.0
        converted_diet = 1.0 if i[0][2] == "poor" else 0.0
        converted_tuple = (converted_smoker, i[0][1],converted_diet)
        converted_data_set.append((converted_tuple,i[1]))
    return converted_data_set

class NearestMeanClassifier:
    def train(self, dataset):
        # >>>>> YOUR CODE HERE
        dataset = convert_data(dataset)
        self.mean_points_tuple_list = [] # A list of tuples. Each tuple is in format (mean_tuple,class), there is only one tuple of each class
        # Then we have a dict, where the key is the class and the value is list of four elements as follows: [sum_smoker, sum_age, sum_diet, counter_of_elements]
        # The idea is to iterate just once over all the dataset, and keep summing the features as well as counting them , so we could divied all the sums by the counted elements for each class
        sum_dict = {}
        # Im assuming that the dataset is a list of tuples in format ((smoker, age, diet),class)
        for i in dataset:
            label = i[1] # the class
            list_of_features = [i[0][0], i[0][1], i[0][2], 1] # smoker, age, diet, counter
            current_list = sum_dict.get(label) 
            #If the class is not in the dict we add it, otherwise we sum the items
            if current_list is None:
                sum_dict[label] = list_of_features
            else:
                # We found an element in the dict, so we increase the counter
                current_list[3] = current_list[3] + 1
                # We sum then all the features
                sum_dict[label] = [x + y for x, y in zip(current_list, list_of_features)]
        
        # Now we iterate over the sum_dict, and append the mean tuples , with their class to the mean_points_tuple_list
        for label, summed_list_of_features in sum_dict.items():
            total_of_elements = summed_list_of_features[3]
            
            average_smoker_value = float(summed_list_of_features[0] / total_of_elements)
            average_age_value = float(summed_list_of_features[1] / total_of_elements)
            average_diet_value =  float(summed_list_of_features[2] / total_of_elements)

            self.mean_points_tuple_list.append(((average_smoker_value,average_age_value, average_diet_value), label))
        
        # <<<<< END YOUR CODE

    def predict(self, x):
        # >>>>> YOUR CODE HERE
        #x is in format (str,float,str) we need (float,float,float)
        
        smoker = 1.0 if x[0] == "yes" else 0.0
        diet = 1.0 if x[2] == "poor" else 0.0
        x = (smoker, x[1], diet) 
        
        if not self.mean_points_tuple_list:
            return "You should train first, before predicting"
        #same logic as in the neighbor function
        min_distance_tuple = (float("inf"),None)
        prediction = ""
        #We iterate over all the trainset, and everytime we find a smaller distance we update our tuple
        for i in self.mean_points_tuple_list:
            temp_distance = distance(x, i[0])
            if temp_distance < min_distance_tuple[0]:
                min_distance_tuple = (temp_distance,i)
    
        #we just want the label from the tuple so
        prediction = min_distance_tuple[1][1]
        # <<<<< END YOUR CODE
        return prediction
```

* Build an object of class `NearestMeanClassifier`, train it on the training data, and return the mean vector for each class. You should return a dictionary with keys `less` and `more`. Each key should correspond to a mean vector.

<!-- grade_id: cell-5f7f00ee83c94703, points: 0, grade: false, solution: true -->
```python
clasifier = NearestMeanClassifier()

def build_and_train():
    # >>>>> YOUR CODE HERE
    dataset = gettrain()
    clasifier.train(dataset)
    # We have the data store in the mean_points_tuple_list variable, so we just convert it to a dict and return it
    returnable_dict = {}
    for i in clasifier.mean_points_tuple_list:
        returnable_dict[i[1]] = i[0]
    return returnable_dict
    # <<<<< END YOUR CODE
build_and_train()
```

Output:
```
{'less': (0.17647058823529413, 17.0, 0.11764705882352941),
 'more': (0.3076923076923077, 20.0, 0.3076923076923077)}
```

<!-- grade_id: cell-415891bde4cbde19, points: 15, grade: true, solution: false -->
```python

```

* Predict the test data using the nearest mean classifier and return all test examples for which all three classifiers (decision tree, nearest neighbor and nearest mean) agree. You should return a list containing tuples. Each tuple should contain the datapoint and the prediction i.e.

`[(('no', 50, 'good'), 'less'), ... ]`

<!-- grade_id: cell-f37f3035a32a8f85, points: 0, grade: false, solution: true -->
```python
def predict_test():
    # >>>>> YOUR CODE HERE    
    dataset = gettest()
    dataset_for_training = gettrain()
    
    n_mean_samples = []
    n_neighbor_samples = []
    decision_tree_samples = []
    
    for i in dataset:
        n_mean_samples.append((i,clasifier.predict(i)))
        n_neighbor_samples.append((i, neighbor(i,dataset_for_training)))
        decision_tree_samples.append((i, decision(i)))

    
    list_of_lists = []
    list_of_lists.append(n_mean_samples)
    list_of_lists.append(n_neighbor_samples)
    list_of_lists.append(decision_tree_samples)
    
    agreed_samples = list(set(list_of_lists[0]).intersection(*list_of_lists))
    
    # <<<<< END YOUR CODE
    return agreed_samples
predict_test()
```

Output:
```
[(('no', 23.0, 'good'), 'less'),
 (('no', 15.0, 'poor'), 'more'),
 (('no', 50.0, 'good'), 'less'),
 (('no', 18.0, 'good'), 'less'),
 (('no', 60.0, 'good'), 'less'),
 (('yes', 45.0, 'poor'), 'more')]
```

<!-- grade_id: cell-853c957eaaf81c28, points: 10, grade: true, solution: false -->
```python

```